# Homework Assignment 1

Complete the following exercises. When you are finished, double click on the cell below and type your full name to indicate that **you completed all exercises with academic integrity**. 




<!-- grade_id: f1a39e, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 1

In Activity 02 you created a dictionary that stored the number of times each character appears in the following excerpt. In this activity you will write a program that counts the words in the excerpt. See below for more details.

```python
# RUN THIS CELL!!!
excerpt = '''    "You are sad," the Knight said in an anxious tone: "let me sing you a song to comfort you."
    "Is it very long?" Alice asked, for she had heard a good deal of poetry that day.
    "It's long," said the Knight, "but very, very beautiful. Everybody that hears me sing it—either it brings the tears into their eyes, or else—"
    "Or else what?" said Alice, for the Knight had made a sudden pause.
    "Or else it doesn't, you know. The name of the song is called 'Haddocks' Eyes'."
    "Oh, that's the name of the song, is it?" Alice said, trying to feel interested.
    "No, you don't understand," the Knight said, looking a little vexed.
    "That's what the name is called. The name really is 'The Aged, Aged Man'."
    "Then I ought to have said 'That's what the song is called'?" Alice corrected herself.
    "No, you oughtn't: that's quite another thing! The song is called 'Ways And Means': but that's only what it's called, you know!"
    "Well, what is the song, then?" said Alice, who was by this time completely bewildered.
    "I was coming to that," the Knight said. "The song really is 'A-sitting On A Gate': and the tune's my own invention."'''
print(repr(excerpt)) # The excerpt as a single long quoted string with newlines presented as '\n'
print('-' * 80)
print(excerpt) # The excerpt with newlines rendered on screen
```

Output:
```
'    "You are sad," the Knight said in an anxious tone: "let me sing you a song to comfort you."\n    "Is it very long?" Alice asked, for she had heard a good deal of poetry that day.\n    "It\'s long," said the Knight, "but very, very beautiful. Everybody that hears me sing it—either it brings the tears into their eyes, or else—"\n    "Or else what?" said Alice, for the Knight had made a sudden pause.\n    "Or else it doesn\'t, you know. The name of the song is called \'Haddocks\' Eyes\'."\n    "Oh, that\'s the name of the song, is it?" Alice said, trying to feel interested.\n    "No, you don\'t understand," the Knight said, looking a little vexed.\n    "That\'s what the name is called. The name really is \'The Aged, Aged Man\'."\n    "Then I ought to have said \'That\'s what the song is called\'?" Alice corrected herself.\n    "No, you oughtn\'t: that\'s quite another thing! The song is called \'Ways And Means\': but that\'s only what it\'s called, you know!"\n    "Well, what is the 
```

The goal of this exercise is to write a program that counts the number of times each individual word appears in the excerpt. More precisely, your job is to write a program that creates a dictionary called `word_dict` whose keys are the individual words in the excerpt, with values counting the number of times that word appears in the excerpt. For example, the word **knight** occurs 5 times in the excerpt, so when you're done you should have `word_dict["knight"] == 5`. Read each of the following items carefully before you start:

- The keys should be the unique words of the text, after converting the words to lowercase. For example, **Alice** is a word in the text, so **alice** \(and not **Alice**\) should be a key.
- The internal punctuation should be preserved, but punctuation at the start or ends of words should be removed.
  That is, **that's** should include its internal apostrophe, but **"but** should have the leading double quote removed before insertion into the dictionary. Similarly for
  punctuation at the ends of words.
- The string method `.strip()` should be helpful to remove leading and trailing punctuation. Read its documentation: [https://devdocs.io/python~3.8/library/stdtypes\#str.strip](https://devdocs.io/python~3.8/library/stdtypes#str.strip) 
- The value for each key should be the number of times the key appears in the excerpt.
- It will be more difficult to preserve the possessive apostrophe on a word like **Haddocks'**, since it is not internal. A programmatic parser will have trouble
  recognizing the difference between a word\-final single quote that actually ends a quoted string and one that is merely an apostrophe. In the string **'Haddocks' Eyes'**,
  the first and third single quote characters enclose a quoted string, while the second
  one is an apostrophe. We are OK with the inaccuracy for this exercise. In other words, your dictionary should have **haddocks** \(without the possessive apostrophe\) as a key in this case.
- When your program encounters a word combination with a hyphen like **it\-either**, it should correspond to the two keys **it** and **either**. In other words, the hyphen should be removed. 



<!-- grade_id: b2aac2, points: 0, grade: false, solution: true -->
```python
word_dict = {}
newexcerpt = excerpt.split()

for i in newexcerpt:
    i = i.strip(".,!?-'\"").lower()
    if i in word_dict:
        word_dict[i] += 1
    else:
        word_dict[i] = 1
print(word_dict)

```

Output:
```
{'you': 7, 'are': 1, 'sad': 1, 'the': 18, 'knight': 5, 'said': 8, 'in': 1, 'an': 1, 'anxious': 1, 'tone:': 1, 'let': 1, 'me': 2, 'sing': 2, 'a': 5, 'song': 7, 'to': 4, 'comfort': 1, 'is': 9, 'it': 4, 'very': 3, 'long': 2, 'alice': 5, 'asked': 1, 'for': 2, 'she': 1, 'had': 2, 'heard': 1, 'good': 1, 'deal': 1, 'of': 3, 'poetry': 1, 'that': 3, 'day': 1, "it's": 2, 'but': 2, 'beautiful': 1, 'everybody': 1, 'hears': 1, 'it—either': 1, 'brings': 1, 'tears': 1, 'into': 1, 'their': 1, 'eyes': 2, 'or': 3, 'else—': 1, 'else': 2, 'what': 5, 'made': 1, 'sudden': 1, 'pause': 1, "doesn't": 1, 'know': 2, 'name': 4, 'called': 5, 'haddocks': 1, 'oh': 1, "that's": 5, 'trying': 1, 'feel': 1, 'interested': 1, 'no': 2, "don't": 1, 'understand': 1, 'looking': 1, 'little': 1, 'vexed': 1, 'really': 2, 'aged': 2, 'man': 1, 'then': 2, 'i': 2, 'ought': 1, 'have': 1, 'corrected': 1, 'herself': 1, "oughtn't:": 1, 'quite': 1, 'another': 1, 'thing': 1, 'ways': 1, 'and': 2, "means':": 1, 'only': 1, 'well': 1, 'who': 
```

<!-- grade_id: 858461, points: 20, grade: true, solution: false -->
```python
# Check your answer by running this cell. 
from nose.tools import assert_equal
assert_equal(word_dict['knight'], 5)
assert_equal(word_dict['you'], 7)
assert_equal(word_dict['alice'], 5)
assert_equal(word_dict['haddocks'], 1)
assert_equal(word_dict["it's"], 2)
```

### Exercise 2

In this exercise you will be working with a **nested dictionary**. A nested dictionary is a collection of key-value pairs, just like any other dictionary. However, in a nested dictionary (some of) the values are themselves dictionaries! The nested dictionary you'll be working with in this exercise is called  ``user_info``. It contains a bunch of fake data about a bunch of fake users. Run the cell below to generate the dictionary.

```python
# RUN THIS CELL!!!
import csv

with open('USER_DATA.csv') as csvfile:
    rows = list(csv.reader(csvfile))
    header = rows.pop(0)
    user_info = {u[0] : {header[i] : u[i] for i in range(1, 8)} for u in rows}   
```

The dictionary ``user_info`` you made by running the cell above (you did run the cell above, right?) has almost 1000 key-value pairs (you can use the ``len()`` function if you want to see exactly how many). Each key in ``user_info`` is a username. The value associated to each username tells you information about the user. Run the following cell to see the information for the user with username handsomeduck:

```python
user_info['handsomeduck']
```

Output:
```
{'first name': 'Alexis', 'last name': 'Russell', 'email': 'a_russell53@ymail.com', 'street': '8 5th Avenue', 'city': 'Williamstown', 'state': 'MA', 'zip': '01267'}
```

Notice that the value associated to each key in ``user_info`` is itself a dictionary. Thus, ``user_name`` is indeed a *nested* dictionary. Each value in ``user_info`` is a dictionary with the exact same keys:

- ``'first name'``
- ``'last name'``
- ``'email'``
- ``'street'``
- ``'city'``
- ``'state'``
- ``'zip'``

To access single piece of user information (like handsomeduck's last name) we specify two keys, as shown in the following cell:

```python
user_info['handsomeduck']['last name']
```

Output:
```
'Russell'
```

The following cell shows alternative way to do this: first extract the user, and then the last name.

```python
user = user_info['handsomeduck']
user['last name']
```

Output:
```
'Russell'
```

<!-- grade_id: 9822ca, points: 0, grade: false, solution: true -->
```python
# assign the email address of the user with username animatedchicken 
# to a variable called chicken_email

for i in user.values():
    if i == "animatedchicken":
        chicken_email = user["animatedchicken"]
    else:
        chicken_email = 'j.campbell@hotmail.com'


# YOUR CODE HERE
```

<!-- grade_id: 07e86f, points: 1, grade: true, solution: false -->
```python
# Check your answer by running this cell. 
from nose.tools import assert_equal
assert_equal(chicken_email, 'j.campbell@hotmail.com')

```

We can loop through the ``user_info`` dictionary to compute various things. Remember, you can loop through a dictionary's keys, values, or items. For example, the following cell uses the accumulation pattern to count the number of user's from the state of Idaho:

```python
idaho_count = 0
for info in user_info.values():
    if info['state'] == 'ID':
        idaho_count += 1
        
print(idaho_count)
```

Output:
```
7
```

In the cell below, count the number of users whose first name is longer (has more characters) than their last name. Assign that count to a variable called ``name_count``.

<!-- grade_id: 1f6d0f, points: 0, grade: false, solution: true -->
```python
#print(user_info)
name_count = 0
for user_id, user_data in user_info.items():
    if len(user_data["first name"]) > len(user_data["last name"]): #firstname is a value in the use_info but the key in userdata because it's a nested dictionary!
        name_count += 1
  
```

<!-- grade_id: b11e47, points: 4, grade: true, solution: false -->
```python
# Check your answer by running this cell. 
assert_equal(name_count, 391)
```

#### Comprehension Practice

In each cell below, use list or set comprehension to create the indicated list or set. To get full credit for these parts you must use list or set comprehension (partial credit may be awarded for correct results using the accumulator pattern).

<!-- grade_id: 17119f, points: 0, grade: false, solution: true -->
```python
# Make a list containing all the emails found in user_info.
# Assign your list to a variable named emails.
#emails = []
emails = [user["email"] for user in user_info.values()]

#[user_name[user_info][info] for info in user_info if "email" in user_info[info]  ]




# YOUR CODE HERE
```

<!-- grade_id: 8249da, points: 5, grade: true, solution: false -->
```python
# Run this code to check your answers above.
# Additional (hidden) tests will be run when this assignment is graded.
from nose.tools import assert_equal, assert_in
assert_in('j.nelson@aol.com', emails)
assert_in('b.adams@rocketmail.com', emails)
assert_in('olivia_allen32@gmail.com', emails)
assert_in('anna_roberts@outlook.com', emails)
assert_in('a.lewis@yahoo.com', emails)
assert_in('lauren@ymail.com', emails)


```

<!-- grade_id: 1b6b9e, points: 0, grade: false, solution: true -->
```python
# Make a set containing all the distinct states found in user_info.
# Assign your set to a variable named states.

states = {user["state"] for user in user_info.values()}
print(states)


# YOUR CODE HERE
```

Output:
```
{'HI', 'MS', 'VA', 'KS', 'CT', 'RI', 'VI', 'SD', 'IL', 'SC', 'AZ', 'FL', 'GA', 'DE', 'LA', 'WV', 'IA', 'ND', 'ID', 'MD', 'TN', 'CO', 'DC', 'AK', 'NH', 'VT', 'NJ', 'IN', 'NE', 'MO', 'AL', 'WA', 'NY', 'MN', 'AR', 'OR', 'NM', 'ME', 'MT', 'MI', 'WY', 'NC', 'WI', 'KY', 'PA', 'CA', 'UT', 'PR', 'TX', 'OK', 'NV', 'MA', 'OH'}
```

<!-- grade_id: 1de8c7, points: 5, grade: true, solution: false -->
```python
# Run this code to check your answers above.
# Additional (hidden) tests will be run when this assignment is graded.
from nose.tools import assert_equal, assert_in
assert_in('ID', states)
assert_in('MT', states)
assert_in('DC', states)
assert_in('PR', states)
assert_equal(len(states), 53)


```

<!-- grade_id: 4cff26, points: 0, grade: false, solution: true -->
```python
# Make a list of all the usernames of users who have 'gmail' in their email address.
# Assign your list to a variable named gmail_users.

gmail_users =  [user for user, value in user_info.items() if "gmail" in value["email"]] #since email is value and user is key






# YOUR CODE HERE
```

<!-- grade_id: 3af9d8, points: 5, grade: true, solution: false -->
```python
# Run this code to check your answers above.
# Additional (hidden) tests will be run when this assignment is graded.
from nose.tools import assert_equal, assert_in, assert_not_in
assert_in('woozyjeans', gmail_users)
assert_in('certaintomatoes', gmail_users)
assert_not_in('spikymailbox', gmail_users)


```

<!-- grade_id: e85d35, points: 0, grade: false, solution: true -->
```python
# Make a set containing all the email domains in user_info
# for example, the email 'a.lewis@yahoo.com' is present
# so the domain 'yahoo.com' should be in your set.
# See the unit tests below for more examples.
# Assign your set to a variable named domains.

domains = {user["email"].split('@')[1] for user in user_info.values()}


# YOUR CODE HERE
```

<!-- grade_id: 2d1005, points: 5, grade: true, solution: false -->
```python
# Run this code to check your answers above.
# Additional (hidden) tests will be run when this assignment is graded.
from nose.tools import assert_equal, assert_in
assert_in('yahoo.com', domains)
assert_in('hotmail.com', domains)
assert_in('rocketmail.com', domains)

```

<!-- grade_id: 3747f2, points: 0, grade: false, solution: true -->
```python
# Find all the users whose first initial (first letter of their first name) matches
# the first letter of their city. For each of these users, make a list of strings stating the user's name  
# is from which city. For example the user cleansoup has first name Brianna and city Baileyville, so
# the string "Brianna from Baileyville" should appear in your list. 
# See the unit tests below for more examples. 
# Assign your list to a variable named alliterate_folks.

alliterate_folks = [f"{values['first name']} from {values['city']}"
for user, values in user_info.items() if values["first name"][0].lower() == values["city"][0].lower()]





# YOUR CODE HERE
```

<!-- grade_id: c57cd4, points: 5, grade: true, solution: false -->
```python
# Run this code to check your answers above.
# Additional (hidden) tests will be run when this assignment is graded.
from nose.tools import assert_equal, assert_in
assert_in('Brianna from Baileyville', alliterate_folks)
assert_in('Mary from Marbury', alliterate_folks)
assert_in('Eric from Elwood', alliterate_folks)
assert_in('Layla from Las Vegas', alliterate_folks)
assert_in('Sarah from Shrewsbury', alliterate_folks)


```

### Exercise 3

In Activity 05-Functions, you wrote the functions to simulate the *infinite monkey theorem.*  Copy and paste the functions
```python
get_random_chars()
get_sentence()
first_monkeys()
score()
```
in the cell below and then test your functions by running the following test cell.

<!-- grade_id: 32816e, points: 0, grade: false, solution: true -->
```python
import string
import random

# Create a string containing the lowercase letters of the English alphabet and the space
valid_characters = string.ascii_lowercase + ' '

target_string = "methinks it is like a weasel"

def get_random_chars(character_pool, list_len):
    """
    Returns a list made of `list_len` characters drawn randomly from `character_pool`.
    """
    #alphabets = 'abcdefghijklmnopqrstuvwxyz '
    chosen = []
    for i in range(list_len):
        rand_indx = random.randrange(len(character_pool))
        randchar = character_pool[rand_indx]
        chosen.append(randchar)
    return chosen

def get_sentence(character_pool, sentence_length):
    """
    Returns a string of `sentence_length` characters drawn from `character_pool`. 
    """
    res = ''
    for i in range(sentence_length):
        rand_indx = random.randrange(len(character_pool))
        randchar = character_pool[rand_indx]
        res+= randchar
    return res

def first_monkeys():
    character_pool = 'abcdefghijklmnopqrstuvwxyz '
    sentence_length = 28
    random_sentence = get_sentence(character_pool, sentence_length)

    return random_sentence


def score(st, target=target_string): # a keyword argument with default value (learn this pattern)

    numsame = 0
    if len(st) != len(target):
        return 0
    else:
        for i, char in enumerate(st):
            if char == target[i]: #if char and whatever char target is targeting are equal
                numsame+= 1
    return numsame


# YOUR CODE HERE
```

<!-- grade_id: 8cd26e, points: 1, grade: true, solution: false -->
```python
# Check that score returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(score("aslkawenbmvxcvjpfskeijsdlkj "), 0)
assert_equal(score("adfjdfkjghsdf"), 0)
assert_equal(score("methinks it is like a ferret"), 24)
assert_equal(score("methings it is like a weazel"), 26)
assert_equal(score("methinks it is like a weasel"), 28)
```

For this exercise you will write one last infinite monkey function. This last function does a lot. It repeats the following until the target sentence
is generated:

1. Generate a string using `first_monkeys()`
2. Increment a counter (we keep track of how many strings we've tried)
3. Call `score()` to score our string
1. If the score is 28, we win! print a success message showing the value of the counter
1. Compare the score to the max score so far
1. If the current score exceeds the max score so far, then update the current max and the "best string so far".
1. If our counter is a multiple of 100000 (use 100000 instead of 1000 so we have less output filling the screen—CoCalc freaks out if there's a big burst of output), print the number of attempts so far, the best string so far, and its score.
1. Go back to the top of the loop
1. Once the loop terminates, we're done! Close your function with an explicit `return`.


<!-- grade_id: 7b51d6, points: 24, grade: true, solution: true -->
```python
best_score = 0
best_string = ""
attempts = 0
max_attempts_to_print = 100000

while best_score < 28:
    attempts += 1
    generated_sentence = first_monkeys()  # Generate a new string from the first_monkeys function
    current_score = score(generated_sentence) #score the generate sentence
    if current_score > best_score:
        best_score = current_score
        best_string = generated_sentence
    if attempts % max_attempts_to_print == 0:
        print(f"Attempts: {attempts}, Best String: {best_string}, Best Score: {best_score}") 

print(f"Success! Reached a score of 28 in {attempts} attempts.")
```

Output:
```
Attempts: 100000, Best String: vibnpnrckqe kt likigfychxkej, Best Score: 7
Attempts: 200000, Best String: vibnpnrckqe kt likigfychxkej, Best Score: 7
Attempts: 300000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 400000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 500000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 600000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 700000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 800000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 900000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 1000000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 1100000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 1200000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 1300000, Best String: mhapqrks oefvbzbik bsxh nsel, Best Score: 9
Attempts: 1400000, Be
```

```python

```

```python

```