# Homework Assignment 5

Complete the following exercises. When you are finished, double click on the cell below and type your full name (as a comment) to indicate that **you completed all exercises with academic integrity**. 

<!-- grade_id: 26003e, points: 1, grade: true, solution: true -->
```python
#Rakeb Abraham Jebessa

```

## Exercise 1: Doubly Linked Lists

Every linked list we have seen so far is called a *singly linked list* because each node has a single reference to the next node in the sequence. An alternative implementation is known as a *doubly linked list*. In a doubly linked list, each node has a reference to the next node (commonly called *next*) as well as a reference to the preceding node (commonly called *back*). Each doubly linked list object has two references: one to the first item in the list (the *head*) and one to the last item in the list (the *tail*). The goal of this exercise it to implement a doubly linked list.

### Step 1: The Nodes

The Node class we have been using for singly linked lists is given below. Modify the definition of the class so that the nodes will work for doubly linked lists. In particular, you make the appropriate modifications to include a new property attribute called *back*.

<!-- grade_id: 8e546c, points: 5, grade: true, solution: true -->
```python
class Node:
    """A node of a doubly linked list"""

    def __init__(self, node_data):
        self._data = node_data
        self._next = None
        self.prev = None
        self.back = None

    def get_data(self):
        """Get node data"""
        return self._data

    def set_data(self, node_data):
        """Set node data"""
        self._data = node_data

    data = property(get_data, set_data)

    def get_next(self):
        """Get next node"""
        return self._next

    def set_next(self, node_next):
        """Set next node"""
        self._next = node_next
    next = property(get_next, set_next)
    def get_back(self):
        """Get the previous node"""
        return self.prev
    def set_back(self, node_prev):
        self.prev = node_prev

    back =  property(get_back, set_back)

    def __str__(self):
        """String"""
        return str(self._data)
```

<!-- grade_id: 96bad3, points: 5, grade: true, solution: false -->
```python
# Unit Tests for the Node class
from nose.tools import assert_equal, assert_true
a = Node(1)
b = Node('two')
c = Node(3)
a.next = b
b.back = a
b.next = c
c.back = b

assert_equal(a.next.data, 'two')
assert_true(a.back is None)
assert_equal(b.back.data, 1)
assert_equal(b.next.data, 3)
assert_equal(c.back.data, 'two')
assert_equal(c.back.back.data, 1)
```

### Step 2: Implement DLinkedList

Write a class for doubly linked lists called `DLinkedList`. A `DLinkedList` object should have two instance variables called `head` and `tail` respectively.  These instance variables should reference the appropriate nodes in the doubly linked list. The class should have all the same methods as the `UnorderedList` class, but you'll need to modify some of their implementations to appropriately incorporate the new features (tails and backs). Some of the methods can use the same code from `UnorderedList`. These methods are provided below. To complete this step you need to include the new instance variable in the definition of the constructor, and write the definitions of the following:

1. `add(item)`
2. `remove(item)`
3. `insert(pos, item)`
4. `append(item)`
5. `pop()` and `pop(pos)`

Write your methods to be as efficient as possible. In particular, your **methods `append(item)` and `pop()` should have constant runtime** (in other words, they should run in O(1)). The other methods, including `pop(pos)`, do not need to have constant runtime. 

There are unit tests for each of the 5 methods in the subsequent cells. Those unit tests should tell you if the methods are functioning properly, but they won't tell you if your algorithms have efficient runtimes. You need to check that on your own!


<!-- grade_id: 11d3af, points: 10, grade: true, solution: true -->
```python
class DLinkedList:
    def __init__(self):
        self.head = None
        self.tail = None

    def is_empty(self):
        return self.head == None
    
    def size(self):
        current = self.head
        count = 0
        while current is not None:
            count = count + 1
            current = current.next

        return count
    
    def search(self, item):
        current = self.head
        while current is not None:
            if current.data == item:
                return True
            current = current.next

        return False
    
    def index(self, item):
        pos = 0
        current = self.head
        while current is not None and current.data != item:
            current = current.next
            pos += 1
        if current is None:
            raise Exception(f"{item} is not in list")
        return pos
    
    #Implement the add method here:
        
    def add(self, item):
        new_node = Node(item)

        if self.head is None:
            self.head = new_node
            self.tail = new_node
        else:
            new_node.next = self.head
            self.head.back = new_node
            self.head = new_node


    #Implement the remove method here:
    def remove(self, item):
        current = self.head
        while current:
            if current.data == item:
                if current.back:
                    current.back.next = current.next
                else:
                    self.head = current.next

                if current.next:
                    current.next.back = current.back
                else:
                    self.tail = current.back

                return

            current = current.next

        print(f"{item} not found in the list")





    #Implement the insert method here:
    def insert(self, pos, item):
        if pos < 0 or pos > self.size():
            raise IndexError("Index out of bounds")
        else:
            new_node = Node(item)

            if self.head == None:
                # If the list is empty, set both head and tail to the new node.
                self.head = new_node
                self.head.back = None
                self.tail = new_node
            elif pos == 0:
                # Insert at the beginning.
                new_node.next = self.head
                self.head.back = new_node
                self.head = new_node
                self.head.back = None
            elif pos == self.size():
                # Insert at the end.
                new_node.back = self.tail
                self.tail.next = new_node
                self.tail = new_node
                new_node.next = None
            else:
                # Insert in the middle.
                current = self.head
                count = 0
                while current is not None:
                    if count == pos-1:
                        break
                    current = current.next
                    count+=1
                
                new_node.back = current
                new_node.next = current.next
                current.next.back = new_node
                current.next = new_node
    

    def append(self, item):
        new_node = Node(item)
        if self.head is None:
            self.head = new_node
            self.tail = new_node
        else:
            new_node.back = self.tail
            self.tail.next = new_node
            self.tail = new_node
        return
    # Implement the pop method here.
    # Make sure pop() has constant runtime.
    # pop(pos) does not need to have constant runtime.

    # ... (previously defined methods and attributes)

    def pop(self, pos=None):
        if self.head is None:
            print("List is empty")
            return None

        if pos is None:
            # If pos is not provided, pop the last element (tail)
            if self.tail is not None:
                data = self.tail.data
                if self.tail.back is not None:
                    self.tail.back.next = None
                self.tail = self.tail.back
                return data
        else:
            # If pos is provided, pop the element at the specified position
            if pos == 0:
                # Pop the head element
                data = self.head.data
                self.head = self.head.next
                if self.head is not None:
                    self.head.back = None
                return data
            else:
                current = self.head
                for i in range(pos):
                    if current is None:
                        print(f"Position {pos} out of range")
                        return None
                    current = current.next

                if current is not None:
                    data = current.data
                    if current.back is not None:
                        current.back.next = current.next
                    if current.next is not None:
                        current.next.back = current.back
                    return data

        print("Position out of range")
        return None
```

<!-- grade_id: bbe40a, points: 5, grade: true, solution: false -->
```python
# Unit Tests for add
from nose.tools import assert_equal
dllst = DLinkedList()
dllst.add(12)
dllst.add(23)
dllst.add(34)
dllst.add(45)

assert_equal(dllst.size(), 4)

node = dllst.tail
for d in [12, 23, 34, 45]:
    assert_equal(node.data, d)
    node = node.back
```

<!-- grade_id: c6622f, points: 5, grade: true, solution: false -->
```python
# Unit Tests for remove
from nose.tools import assert_equal
dllst = DLinkedList()
dllst.add('a')
dllst.add(12)
dllst.add(23)
dllst.add('b')
dllst.add(34)
dllst.add(45)
dllst.add('c')

dllst.remove('b')

node = dllst.tail
for d in ['a', 12, 23, 34, 45, 'c']:
    assert_equal(node.data, d)
    node = node.back
    
dllst.remove('a')

node = dllst.tail
for d in [12, 23, 34, 45, 'c']:
    assert_equal(node.data, d)
    node = node.back
    
dllst.remove('c')

node = dllst.tail
for d in [12, 23, 34, 45]:
    assert_equal(node.data, d)
    node = node.back
```

<!-- grade_id: 8e8357, points: 5, grade: true, solution: false -->
```python
# Unit Tests for insert
from nose.tools import assert_equal, assert_not_equal
dllst = DLinkedList()
dllst.insert(0, 'o')
assert_equal(dllst.size(), 1)
dllst.insert(1, 'e')
dllst.insert(0, 'd')
dllst.insert(2, 'n')
dllst.insert(4, '!')

node = dllst.tail
for ch in '!enod':
    assert_equal(node.data, ch)
    node = node.back

try:
    dllst.insert(6, '!')
except IndexError as e:
    message = str(e)
assert_not_equal(message, "")
```

<!-- grade_id: 51a9f8, points: 5, grade: true, solution: false -->
```python
# Unit Tests for append
# These Tests will NOT check if append has constant runtime. You need to CHECK THIS YOURSELF
from nose.tools import assert_equal
dllst = DLinkedList()
dllst.add(12)
dllst.add(23)
dllst.add(34)
dllst.add(45)
dllst.append("last")

assert_equal(dllst.size(), 5)

node = dllst.tail
for d in ["last", 12, 23, 34, 45]:
    assert_equal(node.data, d)
    node = node.back


# Testing appending to an empty list
lst2 = DLinkedList()
lst2.append(4)
assert_equal(lst2.size(), 1)
assert_equal(lst2.head.data, 4)
assert_equal(lst2.tail.data, 4)
```

<!-- grade_id: d46807, points: 5, grade: true, solution: false -->
```python
# Unit Tests for pop
# These Tests will NOT check if pop() has constant runtime. You need to CHECK THIS YOURSELF
from nose.tools import assert_equal
dllst = DLinkedList()
dllst.add('a')
dllst.add(12)
dllst.add(23)
dllst.add(34)
dllst.add('b')
dllst.add(45)
dllst.add('c')

x = dllst.pop(2)
assert_equal(x, 'b')

node = dllst.tail
for d in ['a', 12, 23, 34, 45, 'c']:
    assert_equal(node.data, d)
    node = node.back
    
x = dllst.pop()
assert_equal(x, 'a')

node = dllst.tail
for d in [12, 23, 34, 45, 'c']:
    assert_equal(node.data, d)
    node = node.back
    
x = dllst.pop(0)
assert_equal(x, 'c')

node = dllst.tail
for d in [12, 23, 34, 45]:
    assert_equal(node.data, d)
    node = node.back
```

## Exercise 2: Recursive GCD Functions

Recall in Activity 7 Fractions, we used the iterative function below to compute the greatest common divisor (GCD) of two integers.

```python
def gcd(m, n):
    """
    Returns the greatest common divisor of m and n.
    Precondition: n is positive
    """
    while m % n != 0:
        m, n = n, m % n
    return n
```

### First Recursive Version

This iterative function above is based on Euclid's algorithm, which observes that when `m % n` does not equal 0, `GCD(m, n) = GCD(n, m % n)`.  Use **Euclid's observation** to write a recursive version of the gcd function below.

<!-- grade_id: 8aea66, points: 7, grade: true, solution: true -->
```python
def gcd_r(m, n):
    """
    Recursive function which returns the greatest common divisor of m and n.
    Precondition: n is positive
    """
    m = abs(m)
    n = abs(n)
    if n == 0:
        return m
    else:
        return gcd_r(n,  m % n)
```

<!-- grade_id: 4d38ac, points: 5, grade: true, solution: false -->
```python
# Unit test for gcd_r
# These tests will not check if your function is RECURSIVE nor if the function USES Euclid's algorithm. You need to TEST THESE YOURSELF!
from nose.tools import assert_equal
assert_equal(gcd_r(-10,5), 5)
assert_equal(gcd_r(-15,20), 5)
assert_equal(gcd_r(46, 12), 2)
assert_equal(gcd_r(12, 42), 6)
assert_equal(gcd_r(10, 10), 10)
assert_equal(gcd_r(17, 33), 1)
assert_equal(gcd_r(2*5*13*7**2, 2**3*3**2*5**3*7**7), 490)
```

### Second Recursive Version: Euclid's Subtraction Algorithm

Euclid was an ancient Greek mathematician who indeed wrote an algorithm for finding the GCD of two positive integers in his famous book [Euclid's Elements](https://en.wikipedia.org/wiki/Euclid%27s_Elements). However, nowhere in Euclid's Elements will you find a reference to the modern mod operator (%). A more direct translation of Euclid's Algorithm would use the observation that when `m>n` we have `GCD(m, n) = GCD(m-n,n)`. Use this observation (without using the mod operator) to write a second recursive version of the gcd function in the cell below:

<!-- grade_id: f7fd8c, points: 7, grade: true, solution: true -->
```python
def gcd_s(m, n):
    """
    Recursive function which uses subtraction (and not %) to return the greatest common divisor of m and n.
    Precondition: m and n are both positive integers
    """
    if m == n:
        return m
    elif m > n:
        return gcd_s(m - n, n)
    else:
        return gcd_s(m, n - m)
```

<!-- grade_id: 63e7dc, points: 5, grade: true, solution: false -->
```python
# Unit test for gcd_s
# These tests will not check if your function is RECURSIVE nor if the function USES Euclid's SUBTRACTION algorithm. You need to TEST THESE YOURSELF!
from nose.tools import assert_equal
assert_equal(gcd_s(46, 12), 2)
assert_equal(gcd_s(12, 42), 6)
assert_equal(gcd_s(10, 10), 10)
assert_equal(gcd_s(17, 33), 1)
assert_equal(gcd_s(12034, 5241632), 22)
```

```python

```