# Homework Assignment 2

Complete the following exercises. When you are finished, double click on the cell below and type your full name to indicate that **you completed all exercises with academic integrity**. 

<!-- grade_id: 1c689d, points: 1, grade: true, solution: true -->
```python
#Rakeb Abraham Jebessa
```

## Exercise 1: More with the Fraction Class

In Activity 07 you developed the `Fraction` class with the following attributes and methods:

Attributes:
* num     
* den  

Methods:
* init(top,bottom)
* str()
* eq(other_fraction)
* lt(other_fraction)
* le(other_fraction)
* add(other_fraction)
* sub(other_fraction)
* mul(other_fraction)
* truediv(other_fraction)
* reduce()





In this Exercise you will continue developing the `Fraction` class. First, in the next cell, copy and paste your working implementation of the `Fraction` class from Activity 07-Exercises.  The helper function `gcd(m,n)` has been copied for you.  Just as in the activity, you will make changes to the `Fraction` class in this cell and run the updated cell each time prior to testing using the Autograder tests.

<!-- grade_id: 8ca294, points: 0, grade: false, solution: true -->
```python
def gcd(m, n):
    """
    Returns the greatest common divisor of m and n.
    Precondition: n is positive
    """
    while m % n != 0:
        m, n = n, m % n
    return n

class Fraction:
    """
    A class for rational numbers, represented exactly as fractions.
    """
    def __init__(self, top=1, bottom=1):
        """Constructor definition.
        top, bottom are integers"""
        
        if type(top) and type(bottom) is not int:
            raise TypeError("Fractions require integer values")

        elif bottom == 0:
            raise ZeroDivisionError("Fraction denominators can not be zero")
       
        else:
            if bottom < 0: # correct negative denominator
                bottom *= -1
                top *= -1
            self.num = top
            self.den = bottom


    def __str__(self):
        return f"{self.num}/{self.den}"

    def __eq__(self, other_fraction):
        first_num = self.num * other_fraction.den
        second_num = other_fraction.num * self.den

        return first_num == second_num

    ### Implement __lt__ here: less than
    def __lt__(self, other_fraction):
        first_num = self.num * other_fraction.den
        second_num = other_fraction.num * self.den
        
        return first_num < second_num

    ### Implement __le__ here: #less than or equal
    def __le__(self, other_fraction):
        first_num = self.num * other_fraction.den
        second_num = other_fraction.num * self.den

        return first_num <= second_num


    def __add__(self, other_fraction):
        """Add self and other_fraction"""
        new_num = self.num * other_fraction.den + self.den * other_fraction.num
        new_den = self.den * other_fraction.den
        return Fraction(new_num, new_den)

    def __sub__(self, other_fraction):
        new_num = self.num * other_fraction.den - self.den * other_fraction.num
        new_den = self.den * other_fraction.den
        return Fraction(new_num, new_den)

    def __mul__(self, other_fraction):
        new_num = self.num * other_fraction.num
        new_den = self.den * other_fraction.den
        return Fraction(new_num, new_den)

    def __truediv__(self, other_fraction):
        new_num = self.num * other_fraction.den
        new_den = self.den * other_fraction.num
        return Fraction(new_num, new_den)

    def __ne__(self, other_fraction):
        first_num = self.num * other_fraction.den
        second_num = other_fraction.num * self.den
        return first_num != second_num

    def __gt__(self, other_fraction):
        first_num = self.num * other_fraction.den
        second_num = other_fraction.num * self.den
        return first_num > second_num

    def __ge__(self, other_fraction):
        first_num = self.num * other_fraction.den
        second_num = other_fraction.num * self.den

        return first_num >= second_num

    def reduce(self):
        """Allow a Fraction object to reduce itself using gcd()"""
        common = gcd(self.num, self.den)
        self.num = self.num // common
        self.den = self.den // common
    def get_num(self):
        return self.num
    def get_den(self):
        return self.den
    def __pow__(self, p):
        if p >=0:
            return Fraction(self.num ** p, self.den **p)
        else:
            return Fraction (self.den ** abs(p) , self.num ** abs(p))
    def __float__(self):
        return self.num/self.den
    def __repr__(self):
        return f"Fraction({self.num}, {self.den})"
    
        
    

```

Check that your implementation is working so far by running the next cell.

<!-- grade_id: 606fde, points: 1, grade: true, solution: false -->
```python
# Check that fraction operations are correct so far:
from nose.tools import assert_equal
assert_equal(Fraction(-1,7) + Fraction(1,11), Fraction(-4, 77))
assert_equal(Fraction(-1,7) - Fraction(1,11), Fraction(-18, 77))
assert_equal(Fraction(-1,7) * Fraction(1,11), Fraction(-1, 77))
assert_equal(Fraction(-1,7) / Fraction(1,11), Fraction(-11, 7))
assert_equal(Fraction(-1,7) == Fraction(1,11), False)
assert_equal(Fraction(-1,7) == Fraction(-2,14), True)
assert_equal(Fraction(-1,7) < Fraction(1,11), True)
assert_equal(Fraction(-1,7) <= Fraction(1,11), True)
assert_equal(Fraction(-1,7) <= Fraction(-2,14), True)
assert_equal(Fraction(1,7) < Fraction(1,11), False)
```

### Warmup: comparison operators

Complete the comparison operators by adding the methods `__ne__`,   `__gt__()` and `__ge__()` for the not-equal-to, greater-than and greater-than-or-equal-to to the `Fraction` class definition. Again, these are boolean operators, so they should return either `True` or `False`, not a `Fraction`. Note that `f1 != f2` calls `f1.__ne__(f2)`. And use the fact that (as long as denominators are positive) $\frac{a}{b}>\frac{c}{d}$ is equivalent to $ad > bc$.  Run the `Fraction` class definition cell again before testing in the below cell.

<!-- grade_id: 8fef16, points: 6, grade: true, solution: false -->
```python
# Check that fraction operations are correct so far:
assert_equal(Fraction(-1,7) > Fraction(1,11), False)
assert_equal(Fraction(-1,7) >= Fraction(-2,14), True)
assert_equal(Fraction(-1,7) != Fraction(1,11), True)
assert_equal(Fraction(-1,7) >= Fraction(1,11), False)
assert_equal(Fraction(-1,7) != Fraction(-2,14), False)
assert_equal(Fraction(1,7) > Fraction(1,11), True)
```

### Getters

In your above `Fraction` class definition, implement the simple methods `get_num` and `get_den` that will return the numerator and denominator of a fraction.  Run the `Fraction` class definition cell again before testing in the below cell.

<!-- grade_id: 8a17a5, points: 2, grade: true, solution: false -->
```python
# Check the getters are correct:
assert_equal(Fraction(-1,7).get_num(), -1)
assert_equal(Fraction(-1,7).get_den(), 7)
```

### Powers

Add the method `__pow__(p)` where `f1**3` invokes `f1.__pow__(3)` and cubes the fraction `f1`.  And recall what `f1**-1` should produce.

<!-- grade_id: 21ce42, points: 5, grade: true, solution: false -->
```python
# Check that __pow__() is  correct:
assert_equal(Fraction(-1,7)**2, Fraction(1,49))
assert_equal(Fraction(-2,7)**3, Fraction(-8,343))
assert_equal(Fraction(-1,7)**-1, Fraction(-7,1))
assert_equal(Fraction(-2,7)**(-2), Fraction(49,4))
```

### Float conversion

Add the method `__float__()` where `float(f1)` invokes `f1.__float__()` and returns the decimal value of `f1`.

<!-- grade_id: 48d45a, points: 5, grade: true, solution: false -->
```python
# Check that __float__() is  correct:
assert_equal(float(Fraction(-1,2)), -0.5)
assert_equal(float(Fraction(2,5)), 0.4)
             
```

### Default Fraction

Modify the `Fraction` class constructor so that the default fraction is $\frac{1}{1}$.

<!-- grade_id: 0f3caa, points: 5, grade: true, solution: false -->
```python
# Check that the constructor modification is  correct:
assert_equal(Fraction(), Fraction(1,1))
```

### The formal string representation

Every Python object has two different string representations: a *formal* string representation and an *informal* string representation. The function `str()` that we all know and love returns the *informal* string representation. To see the formal string representation you can use the function `repr()`. For lots of objects the formal and informal string representations are identical. However, you already know examples where these functions return different strings. Run the following cell to see that `str("cat")` and `repr("cat")` return different values:

```python
# Run this cell!
print(str("cat"))
print(repr("cat"))

import datetime

time = datetime.datetime.now()
print(str(time))
print(repr(time))
```

Output:
```
cat
'cat'
2023-09-25 05:44:42.274659
datetime.datetime(2023, 9, 25, 5, 44, 42, 274659)
```

The expression `str("cat")` returns the string `"cat"` which has 3 characters `c`, `a`, `t`. On the other hand, `repr("cat")` returns the string `"'cat'"` which has 5 characters `'`, `c`, `a`, `t`, `'` (the single quotes are characters in the return value). For another example, try running the following:

```python
import datetime

time = datetime.datetime.now()
print(str(time))
print(repr(time))
```


The purpose of the `str()` function is to produce a (sometime informal) reader-friendly string representation of the object. On the other hand, a good rule of practice is to have `repr()` always return the string that you type when you want to instantiate the object (which agrees with the examples above).

The `repr()` function is called behind the scenes at various times when you are working with an object. For example, if we print a list, the result shows the list of formal string representations of objects. You can see this by running the following:

```python
import datetime

time = datetime.datetime.now()

stuff = [1/2, "cat", time]
print(stuff)
```

Because of this, it is a good habit to always implement the `__repr__()` method of a class.

Run the following cell to see the default formal representation of a Fraction:

```python
print(repr(Fraction(13, 8)))
```

Output:
```
Fraction(13, 8)
```

Implement the `__repr__()` method in your `Fraction` class above so that it returns the string you type to intantiate the fraction object. For example, if the object corresponds to the fraction $\frac{13}{8}$ then the `__repr__()` method should return the string `"Fraction(13, 8)"`. When you're done, run the cell above to see the difference.

<!-- grade_id: 6ad74b, points: 5, grade: true, solution: false -->
```python
# Run this cell to check repr() for the Fraction class
from nose.tools import assert_equal
assert_equal(repr(Fraction(13, 5)), "Fraction(13, 5)")
assert_equal(repr(Fraction(22, 11)), "Fraction(22, 11)")

```

### Adding Exceptions

Recall in Activity 04 we considered exceptions.  In Python you may also throw (or raise) an exception if a condition arises.  For example, run the following code.

```python
list = [5,6,7,8,'five', 10]

for item in list:
    if type(item) is not int:
        raise TypeError("Only integers allowed")
    print(item)


```

Output:
```
5
6
7
8
TypeError                                 Traceback (most recent call last)
/tmp/ipykernel_2220/2933208422.py in <cell line: 3>()
      3 for item in list:
      4     if type(item) is not int:
----> 5         raise TypeError("Only integers allowed")
      6     print(item)

TypeError: Only integers allowed
```

Modify the `Fraction` class constructor so that it checks to make sure that the numerator and denominator are both integers. If either is not an integer, the constructor should raise an appropriate exception with exception string "Fractions require integer values".

<!-- grade_id: d5f5a0, points: 5, grade: true, solution: false -->
```python
# Check that the constructor modification is correct:
x = 5
y = 'x'
try:
    Fraction(x,y)
except TypeError as e:
    if str(e) == "Fractions require integer values":
        pass
    else:
        raise AssertionError("Did not provide correct error string message")  
else:
    raise AssertionError("Did not provide correct error")
try:
    Fraction(x,x)
except TypeError:
    raise AssertionError("No error should be thrown")

    
```

Modify the `Fraction` class constructor so that it checks to make sure that the denominator is not zero. If the denominator is zero, the constructor should raise an appropriate exception with with exception string "Fraction denominators can not be zero".

<!-- grade_id: 794f9e, points: 5, grade: true, solution: false -->
```python
# Check that the constructor modification is correct:
x = 5
y = 0
try:
    Fraction(x,y)
except ZeroDivisionError as e:
    if str(e) == "Fraction denominators can not be zero":
        pass
    else:
        raise AssertionError("Did not provide correct error string message")  
else:
    raise AssertionError("No error was thrown")
try:
    Fraction(y,x)
except TypeError or ZeroDivisionError:
    raise AssertionError("No error should be thrown")
```

## Exercise 2: Mediants

In number theory, given two reduced fractions, $\frac{a}{b}$ and $\frac{c}{d}$, the *mediant* of the two fractions is defined to be the fraction $\frac{a+c}{b+d}$, also in reduced form.  For example, the mediant of $\frac{1}{5}$ and $\frac{1}{3}$ is $$\text{mediant}\left(\frac{1}{5}, \frac{1}{3}\right)=\frac{1+1}{5+3} = \frac{2}{8} = \frac{1}{4}.$$

Write a function `mediant(f1,f2)` that returns the mediant of two given fractions `f1` and `f2`.  Note that `mediant` is not a method of the `Fraction` class, but rather an external function whose parameters are both ``Fraction`` objects. Of course, you can (and should) use these ``Fraction`` objects' methods in your implementation of the mediant function.

<!-- grade_id: 8ab390, points: 0, grade: false, solution: true -->
```python
def mediant(f1, f2):  
    """
    Computes the mediant of the fractions f1 and f2.
    Returns the mediant as a reduced fraction. 
    Assumes: f1 and f2 are in reduced form.
    """
    f3 =  Fraction(f1.num + f2.num, f1.den + f2.den)
    f3.reduce()
    return  f3
    
```

<!-- grade_id: e5c3d7, points: 10, grade: true, solution: false -->
```python
# Check that the mediant(f1, f2) is  correct:
assert_equal(mediant(Fraction(1,3), Fraction(1,5)), Fraction(1,4))

f3 = mediant(Fraction(2,3), Fraction(4,5))
assert_equal(f3.get_num(), 3)

```

## Exercise 3: Mixed Fractions

Recall that improper fractions  can be written as *mixed fraction* (AKA a mixed number). For example, the fraction $\frac{13}{5}$ can be written as the mixed fraction $2\frac{3}{5}$. Each mixed fraction has a *whole part* and a *fractional part*. For example, $2\frac{3}{5}$ has whole part 2 and fractional part $\frac{3}{5}$. 

The goal of this exercise is to write a class `MixedFraction` that is a subclass of the `Fraction` class. The **only** difference between a `MixedFraction` and a `Fraction` are their string representations. 

### Implementing repr()

The formal string representation of a `MixedFraction` object is pretty straight forward. If you call `repr(x)` when `x` is, say, a mixed fraction corresponding to the fraction $\frac{13}{5}$, the call should return the string ``"MixedFraction(13, 5)"``.  More generally, the `__repr__` method should return the string that you would type to construct the object.

Implement the `MixedFraction` class with the appropriate `__repr__` method below. Don't forget that `MixedFraction` is a subclass of the `Fraction` class. 

<!-- grade_id: b2aa74, points: 0, grade: false, solution: true -->
```python
# Write your implementation of the MixedFraction class here:


class MixedFraction(Fraction):
     
    def __repr__(self):
        return f"MixedFraction({self.num}, {self.den})"
    
    def __str__(self):
        mf = MixedFraction(self.num, self.den)
        mf.reduce()
        mf_num= mf.get_num()
        mf_den = mf.get_den()
        whole_part = abs(mf_num) // abs(mf_den)

        new_num = abs(mf_num) % abs(mf_den)
        new_den = abs(mf_den)
        frac_part = Fraction(new_num, new_den)
        if new_num == 0 :
            if mf_num < 0 or mf_den  < 0:
                return f"-{whole_part}"
            else:
                return f"{whole_part}"
        elif whole_part == 0:
            if mf_num < 0 or mf_den  < 0:
                return f"-{frac_part}"
            else:
                return f"{frac_part}"

      
        else:
            if mf_num < 0 or mf_den  < 0:
                return f"-{whole_part} {frac_part}"
            else:
                return f"{whole_part} {frac_part}"

print(2%3)        
print(-17//5)
```

Output:
```
2
-4
```

<!-- grade_id: 119580, points: 5, grade: true, solution: false -->
```python
# Run this cell to check repr() for the MixedFraction class
from nose.tools import assert_equal
assert_equal(repr(MixedFraction(13, 5)), "MixedFraction(13, 5)")
assert_equal(repr(MixedFraction(22, 11)), "MixedFraction(22, 11)")

```

```python

```

### Implementing str()

The informal string representation of a `MixedFraction` object should give its whole part and fractional part, separated by a space. 

For example, if  `x = MixedFraction(13, 5)` then `str(x)` should be `2 3/5`. Here are some additional examples and details:

* The mixed fraction of $\frac{10}{7}$ is $1\frac{3}{7}$.
* The fractional part of a mixed fraction should always be reduced. For example, the mixed fraction of $\frac{55}{10}$ should be $5\frac{1}{2}$. 
* If the fractional part is 0, you should only show the whole part. For example, the mixed fraction of $\frac{12}{3}$ should be $4$. Similarly the mixed fraction of $\frac{0}{23}$ is $0$. 
* If the whole part is 0, you should only show the fractional part. For example, the mixed fraction of $\frac{2}{3}$ is itself.
* The mixed fraction of a negative fraction is obtained by placing a negative sign to the left of the mixed fraction of its absolute value. For example, the mixed fraction of $\frac{-15}{4}$ is $-3\frac{3}{4}$. 

Implement the `__str__` method  **in the cell above** where you wrote your `MixedFraction` class. Run the cell below to test your answer.


<!-- grade_id: ea057a, points: 15, grade: true, solution: false -->
```python
# Run this cell to check str() for the MixedFraction class
from nose.tools import assert_equal
assert_equal(str(MixedFraction(13, 5)), "2 3/5")
assert_equal(str(MixedFraction(0, 5)), "0")
assert_equal(str(MixedFraction(22, 10)), "2 1/5")
assert_equal(str(MixedFraction(-13, 5)), "-2 3/5")
assert_equal(str(MixedFraction(-22, 10)), "-2 1/5")
assert_equal(str(MixedFraction(22, 11)), "2")
assert_equal(str(MixedFraction(-22, 11)), "-2")
assert_equal(str(MixedFraction(2, 5)), "2/5")
assert_equal(str(MixedFraction(-2, 5)), "-2/5")
assert_equal(MixedFraction(22, 10).num, 22)
assert_equal(MixedFraction(22, 10).den, 10)
assert_equal(MixedFraction(-14, 7).num, -14)
assert_equal(MixedFraction(-14, 7).den, 7)


```

```python

```

```python

```

```python

```